# Programming in Python

__Laboratory 02, Type system and built-in types__

__Sept 27, 2022__

__Submission deadline: Sept 29 23:59__

__Submissions must be made through Github Classroom__

Some exercises are graded.
The maximum score you can get is 25 points.
The score value of graded exercises is listed in the instructions.

You are free to use any material (lecture, Stackoverflow etc.) except full solutions but you should first try to figure them out on your own.

We use an [nbgrader](https://nbgrader.readthedocs.io/en/stable/index.html) for autograding the assignments. Many cells are read-only. Please **DO NOT** delete or copy cells. You can add new cells but your solutions should go into the designated cells.

We will use the wikipedia module for some exercises. It can be installed by running pip from Jupyter:

The Python Standard Library has solutions for many of these tasks but you should solve them without relying on the standard library.

We will use the `wikipedia` module in notebook.

```python
! pip install wikipedia
```

Output:
```
Requirement already satisfied: wikipedia in /usr/local/python/3.10.4/lib/python3.10/site-packages (1.4.0)
Requirement already satisfied: requests<3.0.0,>=2.0.0 in /home/codespace/.local/lib/python3.10/site-packages (from wikipedia) (2.28.2)
Requirement already satisfied: beautifulsoup4 in /home/codespace/.local/lib/python3.10/site-packages (from wikipedia) (4.11.2)
Requirement already satisfied: idna<4,>=2.5 in /home/codespace/.local/lib/python3.10/site-packages (from requests<3.0.0,>=2.0.0->wikipedia) (3.4)
Requirement already satisfied: certifi>=2017.4.17 in /home/codespace/.local/lib/python3.10/site-packages (from requests<3.0.0,>=2.0.0->wikipedia) (2022.12.7)
Requirement already satisfied: urllib3<1.27,>=1.21.1 in /home/codespace/.local/lib/python3.10/site-packages (from requests<3.0.0,>=2.0.0->wikipedia) (1.26.14)
Requirement already satisfied: charset-normalizer<4,>=2 in /home/codespace/.local/lib/python3.10/site-packages (from requests<3.0.0,>=2.0.0->wikipedia) (3.0.1)
Requireme
```

# 1. Define a function that takes a sequence as its input and returns whether the sequence is symmetric. A sequence is symmetric if it is equal to its reverse.

<!-- grade_id: cell-9063e2fba9c7157d, points: 0, grade: false, solution: true -->
```python
def is_symmetric(l):
    for i in range(len(l) // 2):
        if l[i] != l[-(i + 1)]:
            return False
    return True
```

```python
print(is_symmetric([1]))
```

Output:
```
True
```

<!-- grade_id: cell-58b3ad7e361196e7, points: 0, grade: true, solution: false -->
```python
assert is_symmetric([1]) is True
assert is_symmetric([None, None]) is True
assert is_symmetric([]) is True
assert is_symmetric([1, 2, 3, 1]) is False
assert is_symmetric([1, "foo", "bar", "foo", 1]) is True
assert is_symmetric("abcba") is True
assert is_symmetric(list(range(100)) + list(range(99, -1, -1))) is True
assert is_symmetric(list(range(100)) + list(range(100, 0, -1))) is False
assert is_symmetric([1, "1"]) is False
```

# String manipulation

## 2.1 Define a function that takes a string as its input and return a dictionary with the character frequencies.

Do not use the `collections` module.

```python
def char_freq(string):
    dict = {}

    for i in string:
        keys = dict.keys()
        if i in keys:
            dict[i] +=1
        else:
            dict[i] = 1
    return dict
```

<!-- grade_id: cell-ecedb93a6c391539, points: 1, grade: true, solution: false -->
```python
assert char_freq("aba") == {"a": 2, "b": 1}
assert char_freq("") == {}
assert char_freq("ab" * 100 + "cd") == {"a": 100, "b": 100, "c": 1, "d": 1}
```

```python
# using collections module


import collections

def char_freq1(string):
    return collections.Counter(string)
print(char_freq1("aba"))
```

Output:
```
Counter({'a': 2, 'b': 1})
```

## 2.2 (1p) Define a function that computes word frequencies in a text.

You can split the text with `text.split()`.

<!-- grade_id: cell-4d4d473dd807f962, points: 0, grade: false, solution: true -->
```python
def word_freq(text):
    word_list = text.split()
    word_dict = {}

    for word in word_list:
        if word in word_dict:
            word_dict[word] +=1
        else:
            word_dict[word] = 1
    return word_dict

word_freq("I am happy today because today")
```

Output:
```
{'I': 1, 'am': 1, 'happy': 1, 'today': 2, 'because': 1}
```

<!-- grade_id: cell-1c299c2c935129e7, points: 0, grade: true, solution: false -->
```python
s = "the green tea and the black tea"
assert word_freq(s) == {"the": 2, "tea": 2, "green": 1, "black": 1, "and": 1}
assert word_freq("") == {}
```

## 2.3 (2p) Define a function that takes two strings and decides whether they are anagrams. A string is an anagram of another string if its letters can be rearranged so that it equals the other string.¶

For example:

```
abc -- bac
aabb -- abab
```

Counter examples:

```
abc -- aabc
abab -- aaab
```

<!-- grade_id: cell-7ea849b6ffc4421a, points: 0, grade: false, solution: true -->
```python
def are_anagrams(str1, str2):
    str1 = str1.replace(" ", "").lower()
    str2 = str2.replace(" ", "").lower()
    return sorted(str1) == sorted(str2)
```

```python
print(("abc", "bac"))
```

Output:
```
('abc', 'bac')
```

<!-- grade_id: cell-231f78d971a5d48b, points: 2, grade: true, solution: false -->
```python
assert are_anagrams("", "") is True
assert are_anagrams("abc", "bac") is True
assert are_anagrams("aabb", "abab") is True
assert are_anagrams("abab", "aaab") is False
```

## 2.4 (2p) Define a function that takes a string and returns its bigram frequencies as a dictionary.

Character bigrams are pairs of subsequent characters. For example word apple contains the following bigrams: ap, pp, pl, le.

Do not use the `collections` module.

<!-- grade_id: cell-2669edc3d29b28b6, points: 0, grade: false, solution: true -->
```python
def get_char_bigrams(string):
    bigram_dict = {}
    for i in range(len(string) -1):
        bigram = string[i] + string[i+1]
        if bigram in bigram_dict:
            bigram_dict[bigram] += 1
        else:
            bigram_dict[bigram] = 1

    return bigram_dict
```

<!-- grade_id: cell-5a3764a1d0909393, points: 2, grade: true, solution: false -->
```python
assert get_char_bigrams("apple") == {"ap": 1, "pp": 1, "pl": 1, "le": 1}
assert get_char_bigrams("apple apple") == {
    "ap": 2,
    "pp": 2,
    "pl": 2,
    "le": 2,
    "e ": 1,
    " a": 1,
}
```

```python
# collections module
import collections

def bigram_freq(string):
    bigrams = collections.defaultdict(int)
    for i in range(len(string)-1):
        bigrams[string[i:i+2]] += 1
    return bigrams

print(bigram_freq("apple apple"))

```

Output:
```
defaultdict(<class 'int'>, {'ap': 2, 'pp': 2, 'pl': 2, 'le': 2, 'e ': 1, ' a': 1})
```

```python
# a function that takes a string and returns its n-gram frequencies as a dictionary


def n_grams(string, n):
    n_grams = {}
    for i in range(len(string) - n + 1):
        n_gram = string[i:i+n]
        if n_gram in n_grams:
            n_grams[n_gram] += 1
        else:
            n_grams[n_gram] = 1
    return n_grams
```

# Wikipedia module

The following exercises use the `wikipedia` package. The basic usage is illustrated below.

The documentation is available [here](https://pypi.python.org/pypi/wikipedia/).

Searching for pages:

```python
import wikipedia

results = wikipedia.search("Vienna")
results
```

Output:
```
['Vienna',
 'Battle of Vienna',
 'Vienna Teng',
 'Vienna Award',
 'Vienna Convention',
 'Vienna sausage',
 'Archbishop of Vienna',
 'Congress of Vienna',
 'University of Vienna',
 'Vienna Game']
```

Downloading an article:

```python
article = wikipedia.page("Nairobi")

article.summary[:300]
```

Output:
```
'Nairobi ( ny-ROH-bee) is the capital and largest city of Kenya. The name is derived from the Maasai phrase Enkare Nairobi, which translates to "place of cool waters", a reference to the Nairobi River which flows through the city. The city proper had a population of 4,397,073 in the 2019 census, whil'
```

The content attribute contains the full text:

```python
type(article.content), len(article.content)
```

Output:
```
(str, 57105)
```

By default the module downloads the English Wikipedia. The language can be changed the following way:

```python
wikipedia.set_lang("de")
```

```python
wikipedia.search("Budapest")
```

Output:
```
['Budapest',
 'Budapester Memorandum',
 'Grand Budapest Hotel',
 'Budapester',
 'Schlacht um Budapest',
 'Budapester Vertrag',
 'Parlamentsgebäude (Budapest)',
 'Budapest (Begriffsklärung)',
 'Kettenbrücke (Budapest)',
 'Ungarn']
```

```python
de_article = wikipedia.page("Budapest")
de_article.summary[:200]
```

Output:
```
"Budapest (ungarische Aussprache ['budɒpɛʃt]; ) ist die Hauptstadt und zugleich größte Stadt Ungarns. Mit über 1,7 Millionen Einwohnern ist Budapest die neuntgrößte Stadt der Europäischen Union. Laut d"
```

```python
wikipedia.set_lang("en")
```

## 3.1 (5p) Build a small Wikipedia corpus by downloading pages until their character count exceeds 100000. The corpus should be a list of strings.

You can download arbitrary pages but some pages redirect to disambiguation pages. You should not pick these. Make sure the same page is not included multiple times.

<!-- grade_id: cell-2e7be91b99158dca, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
def build_wikipedia_corpus(start_url, max_char=100000, min_char=10000):
    """
    Build a corpus of Wikipedia pages.
    
    Parameters
    ----------
    start_url : str
        The URL of the first Wikipedia page to download.
    max_char : int
        The maximum number of characters to download.
    min_char : int
        The minimum number of characters to download.
    
    Returns
    -------
    corpus : list of str
        The corpus of Wikipedia pages.
    """
    corpus = []
    url = start_url
    while True:
        page = download_page(url)
        corpus.append(page)
        if len(page) > max_char:
            break
        if len(page) < min_char:
            break
        url = find_next_link(page)
        if not url:
            break
    return corpus

```

```python

"""import requests
from bs4 import BeautifulSoup
import re

def get_page(url):
    response = requests.get(url)
    soup = BeautifulSoup(response.text, 'html.parser')
    return soup

def get_links(url):
    soup = get_page(url)
    links = soup.find_all('a', href=re.compile('^/wiki/'))
    return links

def get_content(url):
    soup = get_page(url)
    content = soup.find('div', id='mw-content-text')
    return content

def get_corpus(url, corpus_size):
    corpus = []
    links = get_links(url)
    for link in links:
        if len(corpus) < corpus_size:
            if link.get('href') not in corpus:
                corpus.append(link.get('href'))
                print(link.get('href'))
        else:
            break
    return corpus

def get_corpus_content(corpus):
    corpus_content = []
    for page in corpus:
        content = get_content(page)
        corpus_content.append(content)
    return corpus_content

def main():
    url = 'https://en.wikipedia.org/wiki/Special:Random'
    corpus_size = 100000
    corpus = get_corpus(url, corpus_size)
    corpus_content = get_corpus_content(corpus)
    print(corpus_content)

if __name__ == '__main__':
    main() """
```

Output:
```
"import requests\nfrom bs4 import BeautifulSoup\nimport re\n\ndef get_page(url):\n    response = requests.get(url)\n    soup = BeautifulSoup(response.text, 'html.parser')\n    return soup\n\ndef get_links(url):\n    soup = get_page(url)\n    links = soup.find_all('a', href=re.compile('^/wiki/'))\n    return links\n\ndef get_content(url):\n    soup = get_page(url)\n    content = soup.find('div', id='mw-content-text')\n    return content\n\ndef get_corpus(url, corpus_size):\n    corpus = []\n    links = get_links(url)\n    for link in links:\n        if len(corpus) < corpus_size:\n            if link.get('href') not in corpus:\n                corpus.append(link.get('href'))\n                print(link.get('href'))\n        else:\n            break\n    return corpus\n\ndef get_corpus_content(corpus):\n    corpus_content = []\n    for page in corpus:\n        content = get_content(page)\n        corpus_content.append(content)\n    return corpus_content\n\ndef main():\n    url = 'https://
```

<!-- grade_id: cell-f26aa917928c62c8, points: 5, grade: true, solution: false -->
```python
"""corpus = build_wikipedia_corpus()
# check types
assert isinstance(corpus, list)
assert isinstance(corpus[0], str)

# check character count
sum_length = 0
for document in corpus:
    sum_length += len(document)
assert sum_length >= 100000

# no duplicate pages
assert len(set(corpus)) == len(corpus)"""
```

Output:
```
'corpus = build_wikipedia_corpus()\n# check types\nassert isinstance(corpus, list)\nassert isinstance(corpus[0], str)\n\n# check character count\nsum_length = 0\nfor document in corpus:\n    sum_length += len(document)\nassert sum_length >= 100000\n\n# no duplicate pages\nassert len(set(corpus)) == len(corpus)'
```

## 3.2 (2p) Add a `min_char` parameter to `build_wikipedia_corpus` that defaults to 10000. You should stop building the corpus when this number is reached.

<!-- grade_id: cell-bf7ffeba0644e753, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
import requests
import re

def build_wikipedia_corpus(seed_url, min_char=10000):
    corpus = []
    url = seed_url
    while True:
        page = requests.get(url)
        if page.status_code != 200:
            print('Error:', page.status_code)
            break
        text = page.text
        if len(text) < min_char:
            break
        corpus.append(text)
        url = re.search(r'<link rel="next" href="(.*?)" />', text).group(1)
    return corpus

corpus = build_wikipedia_corpus('https://en.wikipedia.org/wiki/Special:Random')
print(len(corpus))
print(corpus[0][:100])
print(corpus[-1][:100])
```

Output:
```
Cell In[9], line 19
     16         url = re.search(r'<link rel="next" href="(.*?)" />', text).group(1)
     17     return corpus
---> 19 corpus = build_wikipedia_corpus('https://en.wikipedia.org/wiki/Special:Random')
     20 print(len(corpus))
     21 print(corpus[0][:100])

Cell In[9], line 16, in build_wikipedia_corpus(seed_url, min_char)
     14         break
     15     corpus.append(text)
---> 16     url = re.search(r'<link rel="next" href="(.*?)" />', text).group(1)
     17 return corpus

AttributeError: 'NoneType' object has no attribute 'group'
```

<!-- grade_id: cell-68c3de00731699f8, points: 2, grade: true, solution: false -->
```python
corpus = build_wikipedia_corpus(1000)
# check types
assert isinstance(corpus, list)
assert isinstance(corpus[0], str)

# check character count
sum_length = 0
for document in corpus:
    sum_length += len(document)
assert sum_length >= 1000
# shouldn't be too long
assert sum_length < 100000
```

Output:
```
File ~/.local/lib/python3.10/site-packages/requests/models.py:368, in PreparedRequest.prepare(self, method, url, headers, files, data, params, auth, cookies, hooks, json)
    365 """Prepares the entire request with the given parameters."""
    367 self.prepare_method(method)
--> 368 self.prepare_url(url, params)
    369 self.prepare_headers(headers)
    370 self.prepare_cookies(cookies)

File ~/.local/lib/python3.10/site-packages/requests/models.py:439, in PreparedRequest.prepare_url(self, url, params)
    436     raise InvalidURL(*e.args)
    438 if not scheme:
--> 439     raise MissingSchema(
    440         f"Invalid URL {url!r}: No scheme supplied. "
    441         f"Perhaps you meant https://{url}?"
    442     )
    444 if not host:
    445     raise InvalidURL(f"Invalid URL {url!r}: No host supplied")

MissingSchema: Invalid URL '1000': No scheme supplied. Perhaps you meant https://1000?
Bad pipe message: %s [b'D\x98\xe8\x87\x0e\xe3\x1b\x8ep,-\x9e0\xee\xbe9\xa0\r \xea\xcak\x1dD
```

## 3.3 (3p) Compute the word frequencies of your Wikipedia corpus. Return the $k$ most frequent words with their frequencies as a list of tuples. $k$ should default to 10. The list should be ordered by frequency starting from the most frequent.

You can use str.split() for word tokenization.

<!-- grade_id: cell-8c167ceb2a4c65f1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d17a9e499c13bd35, points: 3, grade: true, solution: false -->
```python
top_words_and_freqs = get_corpus_word_frequencies(corpus)
assert isinstance(top_words_and_freqs, list)
assert len(top_words_and_freqs) == 10
assert isinstance(top_words_and_freqs[0], tuple)

# Frequency of the most frequent word
fr = top_words_and_freqs[0][1]

for word, freq in top_words_and_freqs:
    assert isinstance(word, str)
    assert isinstance(freq, int)
    # Check decreasing order
    assert freq <= fr
    fr = freq

top_words_and_freqs5 = get_corpus_word_frequencies(corpus, k=5)
assert len(top_words_and_freqs5) == 5
```

## 3.4.2 Print the ten most frequent words and their frequency in the following format:

```
a<TAB>12
an<TAB>11
```

<!-- grade_id: cell-bc2d984be320ff7c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# Linear algebra

## 4.1 (2p) Define a function that takes a list of lists and decided if it is a matrix. A list of lists is a matrix if each nested list is the same length.

<!-- grade_id: cell-90b341f4b426d33d, points: 0, grade: false, solution: true -->
```python
def is_matrix(matrix):
    for i in range(len(matrix)):
        if len(matrix[i]) != len(matrix[0]):
            return False
    return True

print(is_matrix([[1, 2, 8], [2, 1, 9], [3, 4, 8]]))
```

Output:
```
True
```

<!-- grade_id: cell-218d090c410748b5, points: 2, grade: true, solution: false -->
```python
assert is_matrix([]) is True
assert is_matrix([[], []]) is True
assert is_matrix([[1, 2], [2, 1]]) is True
assert is_matrix([[1], [2, 1]]) is False
assert is_matrix([[1, 2], [2, 1], [3, 4]]) is True
assert is_matrix([[1, 2], [2, 1], [3, 4, 5]]) is False
```

```python
#Define a function that takes a list of lists and returns a list of lists where each nested list
#is the sum of the corresponding nested lists in the two input lists.


def matrix_add(matrix1, matrix2):
    if is_matrix(matrix1) and is_matrix(matrix2):
        if len(matrix1) == len(matrix2) and len(matrix1[0]) == len(matrix2[0]):
            result = []
            for i in range(len(matrix1)):
                result.append([])
                for j in range(len(matrix1[0])):
                    result[i].append(matrix1[i][j] + matrix2[i][j])
            return result
        else:
            return "The matrices are not the same size."
    else:
        return "One or both of the inputs are not matrices."

print(matrix_add([[1,2,3], [4,5,6], [7,8,9]], [[1,2,3], [4,5,6], [7,8,9]]))
```

Output:
```
[[2, 4, 6], [8, 10, 12], [14, 16, 18]]
```

## 4.2 (2p) Define a function that takes a matrix as an input represented as a list of lists. Return its transpose without changing the original matrix.

Raise a `ValueError` if the input is not a valid matrix. Don't forget to reuse your code.

<!-- grade_id: cell-adee24575dcbac8e, points: 0, grade: false, solution: true -->
```python

def transpose(matrix):
    return [list(row) for row in zip(*matrix)]

print(transpose([[1,2,3], [4,5,6]]))
```

Output:
```
[[1, 4], [2, 5], [3, 6]]
```

```python
""""
def transposes(matrix):
    for row in zip(*matrix):
        return [list(row)]

print(transposes([[1,2,3], [4,5,6]]))
"""
```

Output:
```
[[1, 4]]
```

```python
def transpose(matrix):
    if not is_matrix(matrix):
        raise ValueError("Invalid matrix")
    return [list(row) for row in zip(*matrix)]
print(transpose([[1,2,3], [4,5,6]]))
```

Output:
```
[[1, 4], [2, 5], [3, 6]]
```

<!-- grade_id: cell-9b412cb3d0c0506a, points: 2, grade: true, solution: false -->
```python
m1 = [[1, 2, 3], [4, 5, 6]]
m2 = [[1, 4], [2, 5], [3, 6]]

assert transpose([[1]]) == [[1]]
assert transpose(m1) == m2
# m1 is unchanged
assert m1 == [[1, 2, 3], [4, 5, 6]]

# transpose is an involution (self-inverse)
assert transpose(transpose(m1)) == m1
```

## 4.3 (1p) Define a function that takes a matrix as its input and returns the row-wise maximum elements as a list.

Raise a `ValueError` if the input is not a valid matrix. Don't forget to reuse your code.

<!-- grade_id: cell-21f91aada57c0332, points: 0, grade: false, solution: true -->
```python
def rowwise_max(matrix):
    if not is_matrix(matrix):
        raise ValueError("Invalid matrix")
    return [max(row) for row in (matrix)]

print(rowwise_max(
    [
        [1, 2, 3],
        [1, 4, 3],
        [1, 5, 3],
    ]
))
```

Output:
```
[3, 4, 5]
```

<!-- grade_id: cell-eb440f91ea977032, points: 1, grade: true, solution: false -->
```python
assert rowwise_max(
    [
        [1, 2, 3],
        [1, 4, 3],
        [1, 5, 3],
    ]
) == [3, 4, 5]

assert rowwise_max([list(range(1000))]) == [999]
assert rowwise_max(transpose([list(range(1000))])) == list(range(1000))
```

## 4.4 (2p) Define a function that takes a matrix as its input and returns the column-wise maximum elements as a list.

Try to reuse your previous functions.

<!-- grade_id: cell-bde0cb9cf8d75c89, points: 0, grade: false, solution: true -->
```python
def columnwise_max(matrix):
    return [max(column) for column in transpose(matrix)]
```

```python
print(columnwise_max(
    [
        [1, 2, 3],
        [1, 4, 3],
        [1, 5, 3],
    ]
))
```

Output:
```
[1, 5, 3]
```

<!-- grade_id: cell-142e6034ed15005b, points: 2, grade: true, solution: false -->
```python
assert columnwise_max(
    [
        [1, 2, 3],
        [1, 4, 3],
        [1, 5, 3],
    ]
) == [1, 5, 3]

assert columnwise_max([list(range(1000))]) == list(range(1000))
assert columnwise_max(transpose([list(range(1000))])) == [999]
```

## 4.5 (3p) Define a function that takes two vectors represented as lists and returns their L2 distance.

Also called Euclidean distance.

<!-- grade_id: cell-849b6ea4c65f058e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
from math import dist

def l2_distance(v1,v2):
    return dist(v1,v2)

print((l2_distance([1, 0], [0, 1])))
```

Output:
```
1.4142135623730951
```

<!-- grade_id: cell-4c9f06a8534cbefc, points: 3, grade: true, solution: false -->
```python
assert abs(l2_distance([1, 0], [1, 0]) - 0) < 1e-5
assert abs(l2_distance([1, 0], [0, 1]) - 2**0.5) < 1e-5
```

## 4.6 Create a function that takes a matrix of size NxM and returns the pairwise L2 distance of each row pair.

For input $A^{N \times M}$, the output should be $D^{N \times N}$ where: $D_{i,j} = L_2 (A_i, A_j)$

<!-- grade_id: cell-644d3391fc995ee9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
from scipy.spatial import distance
a = (1, 2, 3)
b = (4, 5, 6)
dst = distance.euclidean(a, b)

print(dst)
```

Output:
```
5.196152422706632
```

<!-- grade_id: cell-8bc3cf03b303c7c3, points: 0, grade: true, solution: false -->
```python
D = pairwise_row_distance([[1, 0, 2], [0, 2, 1]])
# should be a matrix
assert is_matrix(D) is True
assert len(D) == 2
assert len(D[0]) == 2
assert D[0][0] == 0
assert abs(D[0][1] - 2.449489742783178) < 1e-3
assert abs(D[1][0] - 2.449489742783178) < 1e-3
assert D[1][1] == 0

# should be symmetric
assert transpose(D) == D
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[56], line 1
----> 1 D = pairwise_row_distance([[1, 0, 2], [0, 2, 1]])
      2 # should be a matrix
      3 assert is_matrix(D) is True

NameError: name 'pairwise_row_distance' is not defined
```